<a href="https://colab.research.google.com/github/caenciso30/Programacion-2022-2/blob/main/01-PythonProgramming/01-IntroProgrammingPython.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Python Programming (very fast) Introduction

## Importing modules
Python has an extensive standard library (https://docs.python.org/3/library/) to perform several task in optimized and easy-to-use way. For basic mathematical operations you can just use the normal operators, but for other standard functions, like `sin`, you could import the standard `math` module, or , much better, use either `numpy` or `scipy` modules
```python
import math as m # imports the math module with an alias m
# import math # imports math module but you need to use math.sin or math.cos
# from math import * # DO NOT do this: exposes all inners from math

x = 5.2
y = m.sin(x)
print(x, y)
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Copy and execute the previous code in the following cell (use shift+enter to execute) 
</div>

<!-- grade_id: cell-1d206d3930590507, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Now replace math by numpy and execute it. Is there any difference?
</div>

<!-- grade_id: cell-d2ac73e7587d20db, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Looking for help
What are the units of the argument for the `sin` function? Are there any other functions provided by the `numpy` module? Check the documentation using either google or, inside a notebook
```
help(np.sin)
help(np)
np.sin?
```

```python
help(np)
```

## Printing and reading to and from the screen

You have already seen how to print to the screen. To read, you can use the function call ```input```. Try it:

```python
var = input("Please write something :\t")
print("You have written : " + var)
print(f"You have written : {var}") # uses a f-string: https://realpython.com/python-f-strings/
```

You can also format the printing string, like you have done in c/c++ . For example, you can write something like

```python
x = 5
y = 2.34
print("x = %04d \t y = %.1f \n y=%20.16e\n"%(x, y, y))
print(f"x = {x:04d} \t y = {y:.1f} \n y={y:20.16e}")
```

Note: if you need to manipulate strings and numbers, you should cast the number to string by using ```str```, as in ```str(x)``` .

<div class="alert alert-block alert-info">
<b>Exercise:</b> Print the sum of two numbers with 5 decimals, filled with 0, and in field of 10 spaces 
</div>

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
x = 32.6
y = -4.7
# YOUR CODE HERE
raise NotImplementedError()
```

## Selection structures : `if`, `if else`

This is a simple example on how to use if to select from several options

```python
m = 5
if m < 1:
    print (str(m) + " is smaller than 1")
elif 1 <= m <= 4:
    print (str(m) + " is in between 1 and 4")
else :
    print (str(m) + " is larger than 4")
```

You can also compare values directly, or even compare several values

```python
print(2 == 3)
print(2 == 2 and 3 > 9)
print([2, 3] == [3, 2])
print([2, 3] == [2, 2])
print([2, 3] == [2, 3])
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Read a number and indicate if the number is even or odd 
</div>

<!-- grade_id: cell-8ccf3fbf15c165e4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Loop structures: for and while

```python
N = 5
i = 0
while i < N:
    print (i, end=" ")
    i = i+1
```

```python
for ii in range(1, 10):
    print(f"{ii}", end=" ")
```

```python
name = "My name" # a string is an array for characters
for c in name:
    print (c, end="-")
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Print values from 0.0, to 1.0 (exclusive), in steps of 0.05 
</div>

<!-- grade_id: cell-a83db8b3e865fef8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<div class="alert alert-block alert-info">
<b>Exercise:</b> Now use numpy to do the same (check for the linspace function) 
</div>

<!-- grade_id: cell-4d62b200c665b8ab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## How to declare functions
Functions are the next construct necessary for the a good program design. A function should be defined by using the keyword ```def``` .

```python
def my_function() : 
    print ("I am the function")
    
my_function() # this calls the function    

```

```python
# this shows how to pass arguments
def my_function(x):
    print ("Value of argument is " + str(x))
    
# my_function() # error, requires argument
my_function(-3)
my_function(0.9)

```

```python
# this shows how to pass arguments, and assign default values
def my_function(x = 0):
    print "Value of argument is " + str(x)
    
my_function() # no longer an error
my_function(-3)
```

```python
# this shows how to pass arguments, assign default values, and also calls by named arguments
def my_function(x = 0, y = 1):
    print "Calling my_function"
    print "x " + str(x)
    print "y " + str(y)
    
my_function() # no longer an error
my_function(-3, 2) # x = -3, y = -2
my_function(y=-3, x=2) # x = 2, y = -3
```

### Exercises

1. Write a function that receives a temperature in F and prints it in C 

<!-- grade_id: cell-6d72db8e857ab57e, points: 0, grade: false, solution: true -->
```python
def FtoC(f):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0764b89cea883781, points: 0, grade: true, solution: false -->
```python
assert FtoC(0) == "The temperature 0 F is equivalent to -17.78 C"
assert FtoC(10.3) == "The temperature 10.3 F is equivalent to -12.06 C"
### BEGIN HIDDEN TEST
assert FtoC(-10.3) == "The temperature -10.3 F is equivalent to -23.50 C"
### END HIDDEN TEST
```

2. The coordinates for the vertexes of a given triangle are written into a file in the form  \begin{align*}x_1\ \ \  y_1\\x_2\ \ \  y_2\\x_3\ \ \ y_3  \end{align*} Write a program that reads this values, and computes the area of the triangle, defined as $$ A = \frac{1}{2}[x_2y_3 - x_3y_2 - x_1y_3 + x_3y_1 + x_1y_2 - x_2y_1]$$. Use functions.

```python

```